---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Let's learn

```yaml
type: VideoExercise
key: 3a249b006a
xp: 50
```

`@projector_key`
f9f93901a0e9d44493e78fef56bb1226

---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}
obj = 1


```

`@sct`
```{python}

```
